# IS318 - Machine Learning
## TP2 - Linear classification

The goal of this TP is to experiment with linear models for classification.

```python
import numpy as np
import matplotlib.pyplot as plt
from sklearn.datasets import make_blobs, make_classification
from sklearn.model_selection import train_test_split
```

### 1. Dataset

We will work with a synthetic binary classification dataset:

```python
X, y = make_blobs(n_samples=500, n_features=2, centers=2, center_box=(-3, 3), random_state=42) # Easy
# X, y = make_classification(n_samples=500, n_features=2, n_redundant=0, n_informative=2) # Harder
plt.scatter(X[:, 0][y == 0], X[:, 1][y == 0])
plt.scatter(X[:, 0][y == 1], X[:, 1][y == 1])
plt.tight_layout()
print(X.shape, y.shape)
```

Before implementing a classification method, we will apply some preprocessing steps.

**(Question)** Preprocessing: normalize the data such that it has zero mean and unit variance. To do that, for each feature, substract the mean and divide by the standard deviation over the samples. Plot the data.

This procedure is called *standardization*.
It is useful for numerical stability (notably for gradient descent).

<!-- grade_id: standardization, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: standardization-tests, points: 1, grade: true, solution: false -->
```python
assert np.allclose(X.mean(axis=0), 0.)
assert np.allclose(X.var(axis=0), 1.)
```

**(Question)** Plot the data points as above to observe the difference of standardization.

<!-- grade_id: cell-c0770150e5c3dd24, points: 1, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

**(Question)** As usual, add a column with ones to simplify dot product notation. Then, split the dataset into training and validation sets.

<!-- grade_id: split, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: split-tests, points: 1, grade: true, solution: false -->
```python
assert X.shape[1] == 3
assert np.all(X[:, 0] == 1.)
```

**(Question)** With real data, it is considered wrong to apply standardization before splitting the dataset. Can you explain why?

### 2. Logistic regression

**(Question)** Complete the `LogisticRegression` class below to implement logistic regression for binary classification.
- For the `fit` method, implement a simple gradient descent on the binary cross entropy loss. Record the evolution of the loss over the iterations in the `loss_hist` attribute.
- For the `predict` method, compute predictions from the model and return the binary decisions by thresholding at $0.5$.

<!-- grade_id: logistic-regression, points: 3, grade: true, solution: true -->
```python
class LogisticRegression():
    
    def __init__(self, learning_rate=0.1, n_iters=1000):
        self.learning_rate = learning_rate
        self.n_iters = n_iters
    
    def sigmoid(self, x):
        # YOUR CODE HERE
        raise NotImplementedError()

    def binary_cross_entropy(self, y, y_pred):
        # YOUR CODE HERE
        raise NotImplementedError()
        
    def fit(self, X, y):
        # YOUR CODE HERE
        raise NotImplementedError()
    
    def predict(self, X):
        # YOUR CODE HERE
        raise NotImplementedError()
```

**(Question)** Apply your logistic regression model with default arguments to our synthetic dataset. Report the classification accuracy on the training and validation sets.

<!-- grade_id: logistic-regression-application, points: 2, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

**(Question)** Plot the decision boundary on top on the data points.

*Hint*: to compute the decision boundary, start from its equation and transform it to an usual affine formulation.

<!-- grade_id: plot-decision-boundary, points: 2, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

**(Question)** Play around with the `learning_rate` and `n_iters` hyperparameters to see the impact on gradient descent convergence.
Plot the evolution of the loss over the iterations for different learning rates.
Comment the results.

<!-- grade_id: cell-e7c41ebfd71dc746, points: 0, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-eeedab9405a61031, points: 0, grade: true, solution: true -->
YOUR ANSWER HERE

### 3. Gaussian generative model

**(Question)** Complete the `GaussianGenerativeModel` class below.

<!-- grade_id: gaussian-generative-model, points: 3, grade: true, solution: true -->
```python
class GaussianGenerativeModel():

    def fit(self, X, y):
        # YOUR CODE HERE
        raise NotImplementedError()
        
    def predict(self, X):
        # YOUR CODE HERE
        raise NotImplementedError()
```

**(Question)** Apply your gaussian generative model to our synthetic dataset. Report the classification accuracy on the training and validation sets.

*Hint:* The column with ones is not needed for this model.

<!-- grade_id: gaussian-generative-model-application, points: 1, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

**(Question)** Plot the decision boundary on top on the data points and compare it with the one obtained with logistic regression.

<!-- grade_id: decision-boundary-lda, points: 0, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

### 4. Comparison with `sklearn`

**(Question)** Apply logistic regression using the `LogisticRegression` class from the `scikit-learn` library. Compare the results with your previous implementation.

*Hints:*
- If your features are already augmented with a column of ones, set the `fit_intercept` parameter to `False`
- Set `penalty` (regularization to `None` to match your unregularized implementation

<!-- grade_id: cell-69a8aa6c12da9bf1, points: 0, grade: true, solution: true -->
```python
from sklearn.linear_model import LogisticRegression as LogisticRegressionSklearn
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-f96b86b29a4215a3, points: 0, grade: true, solution: true -->
YOUR ANSWER HERE

### 5. Application to breast cancer dataset

```python
from sklearn.datasets import load_breast_cancer
dataset = load_breast_cancer()
```

**(Question)** Load the data into `X` and `y` variables. Familiarize yourself with this dataset. What are the dimensions? Explain what they mean.

<!-- grade_id: cell-60ee53c331595ef1, points: 0, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-14189902a0489aeb, points: 0, grade: true, solution: true -->
YOUR ANSWER HERE

**(Question)** Split into train/test and preprocess the data.

<!-- grade_id: breast-cancer-split, points: 1, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

**(Question)** Apply your implementations of logistic regression and gaussian generative model to this dataset.

<!-- grade_id: cell-1d2b82910938cd00, points: 0, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

**(Question)** Apply cross-validation (see `sklearn.model_selection`) on your implementations to successively train several models and obtain predictions for the whole dataset.
Comment the results.

<!-- grade_id: cell-0f3c63ce7aac3838, points: 0, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-5bdf5d2000379a27, points: 0, grade: true, solution: true -->
YOUR ANSWER HERE